// Numbas version: exam_results_page_options
{"name": "Florian's copy of Test your Skills Week 9: Differential Equations 6", "metadata": {"description": "", "licence": "None specified"}, "duration": 0, "percentPass": 0, "showQuestionGroupNames": false, "showstudentname": true, "question_groups": [{"name": "Group", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": [""], "questions": [{"name": "DE6: Critical Damping in RLC circuit", "extensions": [], "custom_part_types": [], "resources": [["question-resources/346px-RLC_series_circuit_v1.svg.png", "/srv/numbas/media/question-resources/346px-RLC_series_circuit_v1.svg.png"], ["question-resources/135px-RLC_series_circuit_v1.svg.png", "/srv/numbas/media/question-resources/135px-RLC_series_circuit_v1.svg.png"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Florian Breuer", "profile_url": "https://editor.newnumbas.net/accounts/profile/28/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<p><img src=\"resources/question-resources/135px-RLC_series_circuit_v1.svg.png\"/></p>\n<p>In a series&nbsp;RLC circuit, the current $I(t)$ as a function of time $t$ satisfies the second-order ordinary differential equation</p>\n<p>\\[<br/>\\frac{d^2}{dt^2}I(t) + \\frac{R}{L}\\frac{d}{dt}I(t) + \\frac{1}{LC}I(t)=0,<br/>\\]</p>\n<p>where&nbsp;</p>\n<ul>\n<li>$R$ is the&nbsp;<em>resistance</em> in ohms ($\\Omega$),</li>\n<li>$L$ is the&nbsp;<em>inductance</em> in henry (H), and</li>\n<li>$C$ is the&nbsp;<em>capacitance</em> in farad (F).</li>\n</ul>\n<p></p>\n<p></p>", "advice": "<p>The characteristic equation for this second-order linear differential equation is</p>\n<p>\\[<br/>r^2 + \\frac{R}{L}r + \\frac{1}{LC} = 0.\\quad\\quad (1)<br/>\\]</p>\n<p>The three types of damping are determined by the type of roots of the characteristic equation.</p>\n<p>The discriminant of this quadratic equation is&nbsp;</p>\n<p>\\[<br/>\\Delta = \\left(\\frac{R}{L}\\right)^2 - 4\\frac{1}{LC} = \\frac{R^2C - 4L}{L^2C}.<br/>\\]</p>\n<ul>\n<li>If (1) has complex roots, then the circuit is <em>under-damped</em>. This happens when $\\Delta &lt; 0$, i.e. $R^2C &lt; 4L$.</li>\n<li>If (1) has distinct real roots, then the circuit is <em>over-damped</em>. This happens when $\\Delta &gt; 0$, i.e. $R^2C &gt; 4L$.</li>\n<li>If (1) has a repeated real root, then the circuit is <em>critically damped</em>. This happens when $\\Delta = 0$, i.e. $R^2C = 4L$.</li>\n</ul>\n<p></p>\n<p></p>\n<p>(a) In our particular example, $R = \\var{R}\\Omega$, $L = \\var{L}$H and $C = \\var{C}$F.&nbsp;<br/>The characteristic equation is&nbsp;</p>\n<p>\\[<br/>\\simplify{r^2 + {R}/{L}*r+1/{L*C}} = 0,<br/>\\]</p>\n<p>which has $\\Delta = \\simplify{({R}^2*{C}-4*{L})/({L}^2*{C})} \\; \\var{ineq} \\; 0$, so the equation has {roots}. Thus, the circuit is {CorrectAnswer}.</p>\n<p></p>\n<p>(b) In order to make the circuit critically damped, we need $R^2C = 4L$, i.e.&nbsp;</p>\n<p>\\[<br/>R = 2\\sqrt{\\frac{L}{C}} = \\simplify{2*sqrt{{L}/{C}}}\\Omega.<br/>\\]</p>\n<p></p>", "rulesets": {}, "variables": {"R": {"name": "R", "group": "Ungrouped variables", "definition": "random(1..9)", "description": "", "templateType": "anything"}, "L": {"name": "L", "group": "Ungrouped variables", "definition": "random(1..9)", "description": "", "templateType": "anything"}, "C": {"name": "C", "group": "Ungrouped variables", "definition": "random(1..9)", "description": "", "templateType": "anything"}, "Answers": {"name": "Answers", "group": "Ungrouped variables", "definition": "[\"critically damped\", \"under-damped\", \"over-damped\"]", "description": "", "templateType": "anything"}, "CorrectAnswer": {"name": "CorrectAnswer", "group": "Ungrouped variables", "definition": "switch (\n  R^2*C < 4*L, Answers[1],\n  R^2*C > 4*L, Answers[2],\n  R^2*C = 4*L, Answers[0]\n  )", "description": "", "templateType": "anything"}, "IncorrectAnswers": {"name": "IncorrectAnswers", "group": "Ungrouped variables", "definition": "switch (\n  R^2*C < 4*L, [Answers[0], Answers[2]],\n  R^2*C > 4*L, [Answers[0], Answers[1]],\n  R^2*C = 4*L, [Answers[1], Answers[2]]\n  )", "description": "", "templateType": "anything"}, "roots": {"name": "roots", "group": "Ungrouped variables", "definition": "switch (\n  R^2*C = 4*L, \"repeated roots\",\n  R^2*C < 4*L, \"complex roots\",\n  R^2*C > 4*L, \"distinct real roots\"\n  )", "description": "", "templateType": "anything"}, "ineq": {"name": "ineq", "group": "Ungrouped variables", "definition": "switch (\n  R^2*C = 4*L, \"=\",\n  R^2*C < 4*L, \"<\",\n  R^2*C > 4*L, \">\"\n  )", "description": "", "templateType": "anything"}}, "variablesTest": {"condition": "4*L <> r^2*C", "maxRuns": 100}, "ungrouped_variables": ["R", "L", "C", "Answers", "CorrectAnswer", "IncorrectAnswers", "roots", "ineq"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "1_n_2", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Suppose that $R = \\var{R}\\Omega$, $L = \\var{L}$H and $C = \\var{C}$F.&nbsp;</p>\n<p>Is the circuit over-damped, under-damped or critically damped?&nbsp;</p>", "minMarks": 0, "maxMarks": 0, "shuffleChoices": true, "displayType": "radiogroup", "displayColumns": 0, "showCellAnswerState": true, "choices": ["{CorrectAnswer}", "{IncorrectAnswers[0]}", "{IncorrectAnswers[1]}"], "matrix": ["1", 0, 0], "distractors": ["", "", ""]}, {"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Given that $L = \\var{L}$H and $C = \\var{C}$F, what value of the resistance $R$ would make this circuit critically damped?</p>\n<p>$R = $&nbsp;[[0]]$\\Omega$</p>\n<p><em>(Either enter an exact answer, e.g. writing sqrt(x) for $\\sqrt{x}$, or give the answer correct to 3 decimal places).</em></p>", "gaps": [{"type": "jme", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "2*sqrt({L}/{C})", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": "0.01", "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "valuegenerators": []}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always"}]}], "allowPrinting": true, "navigation": {"allowregen": true, "reverse": true, "browse": true, "allowsteps": true, "showfrontpage": true, "showresultspage": "oncompletion", "navigatemode": "sequence", "onleave": {"action": "none", "message": ""}, "preventleave": true, "startpassword": ""}, "timing": {"allowPause": true, "timeout": {"action": "none", "message": "<p>Sorry, time's up.</p>"}, "timedwarning": {"action": "none", "message": "<p>5 minute left.</p>"}}, "feedback": {"showactualmark": true, "showtotalmark": true, "showanswerstate": true, "allowrevealanswer": true, "advicethreshold": 0, "intro": "<p>This is Test Your Skills, week 9. It has one question on damped circuits from chapter DE 6.</p>\n<p>Good luck!</p>", "reviewshowscore": true, "reviewshowfeedback": true, "reviewshowexpectedanswer": true, "reviewshowadvice": true, "feedbackmessages": []}, "contributors": [{"name": "Florian Breuer", "profile_url": "https://editor.newnumbas.net/accounts/profile/28/"}], "extensions": [], "custom_part_types": [], "resources": [["question-resources/346px-RLC_series_circuit_v1.svg.png", "/srv/numbas/media/question-resources/346px-RLC_series_circuit_v1.svg.png"], ["question-resources/135px-RLC_series_circuit_v1.svg.png", "/srv/numbas/media/question-resources/135px-RLC_series_circuit_v1.svg.png"]]}