// Numbas version: exam_results_page_options
{"name": "DE5: Non-homogeneous 2nd order ODE (explore)", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "DE5: Non-homogeneous 2nd order ODE (explore)", "tags": [], "metadata": {"description": "<p>Guides the student via a number of steps to the solution of a non-homogeneous second-order constant coefficient linear initial value problem.</p>", "licence": "None specified"}, "statement": "<p>In this question, we will solve the initial value problem</p>\n<p>\\[\\simplify{(y'' + {-r1-r2}y' + {r1*r2}y)} = \\simplify{({c}x + {d})},\\qquad y(0) = \\var{y0},\\quad y'(0) = \\var{y1}.\\]</p>\n<p></p>", "advice": "<p>We start by determining the&nbsp;complementary function by&nbsp;solving the homogeneous version of the differential equation</p>\n<p>\\[\\simplify{(y'' + {-r1-r2}y' + {r1*r2}y)} = 0.\\]</p>\n<p>The characteristic equation is</p>\n<p>\\[\\simplify{(r^2 + {-r1-r2}r + {r1*r2})} = 0,\\]</p>\n<p>and it has two distinct real roots $r_1 = \\var{r1}$ and $r_2 = \\var{r2}$. Thus, the complementary function is</p>\n<p>\\[y_c(x) = C_1e^{\\simplify{({r1}x)}} +&nbsp;C_2e^{\\simplify{({r2}x)}}.\\]</p>\n<p>Next, we must find a particular solution. This should have a similar form to $\\simplify{({c}x+{d})}$, so we try</p>\n<p>\\[y_p(x) = ax+b.\\]</p>\n<p>Since neither of these terms is a term of the complementary function, we're good to go. We plug $y_p(x)$ into the LHS of the differential equation and obtain</p>\n<p>\\[\\simplify{(-{r1}-{r2})*a + {r1}*{r2}(a*x+b)} =&nbsp;\\simplify{({c}x+{d})}.\\]</p>\n<p>Comparing $x$-coefficients and constant coefficients, respecively, we get two linear equations in $a$ and $b$:</p>\n<p>\\begin{align} \\simplify{({r1}*{r2}*a)} &amp;= \\var{c} \\\\ \\simplify{(-{r1}-{r2})*a + {r1}*{r2}*b} &amp;= \\var{d}.\\end{align}</p>\n<p>Solving these, we find that</p>\n<p>\\[a = \\var{a}\\quad\\text{and}\\quad b = \\var{b},\\]</p>\n<p>so&nbsp;</p>\n<p>\\[y_p(x) = \\simplify{({a}x+{b})}\\]</p>\n<p>and the general solution to the differential equation is</p>\n<p>\\[y(x) = y_c(x) + y_p(x) = C_1e^{\\simplify{({r1}x)}} +&nbsp;C_2e^{\\simplify{({r2}x)}} + \\simplify{({a}x+{b})}.\\]</p>\n<p>It remains to solve the initial value problem. We plug in our initial values and obtain the equations</p>\n<p>\\begin{align}\\var{y0} = y(0) &amp;= \\simplify{C_1 + C_2 + {b}} \\\\&nbsp;&nbsp;\\var{y1} = y'(0) &amp;= \\simplify{{r1}*C_1 + {r2}*C_2 + {a}}.\\end{align}</p>\n<p>Solving these, we obtain&nbsp;</p>\n<p>\\[C_1 = \\var{c1}\\quad\\text{and}\\quad C_2 = \\var{c2},\\]</p>\n<p>so our final solution to the initial value problem is</p>\n<p>\\[y(x) = \\simplify{({c1}*e^({r1}x) + {c2}*e^({r2}x) + {a}*x + {b})}.\\]</p>", "rulesets": {}, "extensions": [], "variables": {"r1": {"name": "r1", "group": "Ungrouped variables", "definition": "random(-3..3 except 0)", "description": "<p>One root of characteristic equation</p>", "templateType": "anything"}, "r2": {"name": "r2", "group": "Ungrouped variables", "definition": "random(-3..3 except 0)", "description": "<p>Other root of characteristic equation</p>", "templateType": "anything"}, "a": {"name": "a", "group": "Ungrouped variables", "definition": "random(-5..5 except 0)", "description": "<p>$x$-Coefficient in particular solution $y_p(x) = ax+b$</p>", "templateType": "anything"}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "random(5..5 except 0)", "description": "<p>Constant coefficient in particular solution $y_p(x) = ax+b$</p>", "templateType": "anything"}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "a*r1*r2", "description": "<p>$x$-coefficient of RHS of DE.</p>", "templateType": "anything"}, "d": {"name": "d", "group": "Ungrouped variables", "definition": "-(r1+r2)*a+r1*r2*b", "description": "<p>Constant coeff of RHS of DE</p>", "templateType": "anything"}, "c1": {"name": "c1", "group": "For initial value", "definition": "random(-2..2)", "description": "<p>Value of $c_1$ in final solution $y(x) = c_1e^{r_1x} + c_2e^{r_2x} + ax+b$.</p>", "templateType": "anything"}, "c2": {"name": "c2", "group": "For initial value", "definition": "random(-2..2)", "description": "<p>Value of $c_2$ in final solution $y(x) = c_1e^{r_1x} + c_2e^{r_2x} + ax+b$.</p>", "templateType": "anything"}, "y0": {"name": "y0", "group": "For initial value", "definition": "c1+c2+b", "description": "<p>Initial value of $y(0) = y_0$</p>", "templateType": "anything"}, "y1": {"name": "y1", "group": "For initial value", "definition": "r1*c1+r2*c2+a", "description": "<p>Initial value of $y'(0) = y_1$</p>", "templateType": "anything"}}, "variablesTest": {"condition": "r1 < r2", "maxRuns": 100}, "ungrouped_variables": ["r1", "r2", "a", "b", "c", "d"], "variable_groups": [{"name": "For initial value", "variables": ["c1", "c2", "y0", "y1"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "1_n_2", "useCustomName": true, "customName": "Let's begin!", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Form of complementary function.", "rawLabel": "", "otherPart": 1, "variableReplacements": [], "availabilityCondition": "answered", "penalty": "", "penaltyAmount": 0, "lockAfterLeaving": true}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Off to a good start (bonus point)", "prompt": "<p>This question consists of multiple parts, the end result of which will be a full solution to this initial value problem.&nbsp;</p>\n<p style=\"padding-left: 30px;\">You will have the option to retry some parts, for a limited penalty.</p>\n<p style=\"padding-left: 30px;\">You will sometimes have the option to ask for help. This incurs no penalty.</p>\n<p></p>\n<p>What should we do first?</p>", "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "1", "showCellAnswerState": true, "choices": ["Solve the homogeneous form of the differential equation to find the&nbsp;<em>complementary function.</em>", "Find a particular solution to the non-homogeneous differential equation.", "Plug in the initial values."], "matrix": ["1", "0", 0], "distractors": ["", "This is a dangerous approach, as our first choice of particular solution might just make the LHS = 0.", "We have nothing to plug the initial values into!"]}, {"type": "1_n_2", "useCustomName": true, "customName": "Form of complementary function.", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Complementary function", "rawLabel": "", "otherPart": 2, "variableReplacements": [], "availabilityCondition": "answered", "penalty": "", "penaltyAmount": 0, "lockAfterLeaving": true}, {"label": "Help, please! ", "rawLabel": "Help, please! ", "otherPart": 6, "variableReplacements": [], "availabilityCondition": "not (answered and credit=1)", "penalty": "", "penaltyAmount": 0, "lockAfterLeaving": false}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Complementary function", "prompt": "<p>We first solve the homogeneous form of the differential equation: $\\simplify{(y'' +{-r1-r2}y' + {r1*r2}y)}=0$.&nbsp;</p>\n<p>The solution to the homogeneous differential equation has the form</p>", "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "1", "showCellAnswerState": true, "choices": ["$y_c(x) = C_1e^{r_1x} + C_2e^{r_2x}$.", "$y_c(x) = (C_1x+C_2)e^{rx}$.", "$y_c(x) = e^{\\alpha x}\\big(C_1\\cos(\\beta x)+C_2\\sin(\\beta x)\\big)$"], "matrix": ["1", 0, 0], "distractors": ["", "", ""]}, {"type": "gapfill", "useCustomName": true, "customName": "Complementary function", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Form of particular solution", "rawLabel": "", "otherPart": 3, "variableReplacements": [], "availabilityCondition": "answered and credit=1", "penalty": "", "penaltyAmount": 0, "lockAfterLeaving": true}, {"label": "Try again!", "rawLabel": "Try again!", "otherPart": 2, "variableReplacements": [], "availabilityCondition": "answered and credit<1", "penalty": "Retried complementary function", "penaltyAmount": "0.5", "lockAfterLeaving": false}, {"label": "Well, never mind. Let's continue to the next step.", "rawLabel": "Well, never mind. Let's continue to the next step.", "otherPart": 3, "variableReplacements": [], "availabilityCondition": "answered and credit<1", "penalty": "", "penaltyAmount": 0, "lockAfterLeaving": false}, {"label": "Help, please!", "rawLabel": "Help, please!", "otherPart": 6, "variableReplacements": [], "availabilityCondition": "not (answered and credit=1)", "penalty": "", "penaltyAmount": 0, "lockAfterLeaving": false}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Complementary function", "prompt": "<p>So the complementary function is of the form $y_c(x) = C_1e^{r_1x} + C_2e^{r_2x}$.</p>\n<p>We compute the values of $r_1$ and $r_2$, and obtain</p>\n<p>$r_1 = $ [[0]]&nbsp; &nbsp; and</p>\n<p>$r_2 = $&nbsp;[[1]]</p>\n<p><em>(If you get this wrong, you may retry this part and lose 0.5 marks. Thereafter you may retry this part as often as you like without losing further marks.)</em></p>", "gaps": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": "0.5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minValue": "r1", "maxValue": "r1", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": "0.5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minValue": "r2", "maxValue": "r2", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "sortAnswers": true}, {"type": "1_n_2", "useCustomName": true, "customName": "Form of particular solution", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Particular solution", "rawLabel": "", "otherPart": 4, "variableReplacements": [], "availabilityCondition": "answered", "penalty": "", "penaltyAmount": 0, "lockAfterLeaving": true}, {"label": "Help, please!", "rawLabel": "Help, please!", "otherPart": 7, "variableReplacements": [], "availabilityCondition": "not (answered and credit=1)", "penalty": "", "penaltyAmount": 0, "lockAfterLeaving": false}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Particular solution", "prompt": "<p>Next, we must find a particular solution to the differential equation. The correct form of the particular solution is</p>", "minMarks": 0, "maxMarks": 0, "shuffleChoices": false, "displayType": "radiogroup", "displayColumns": "1", "showCellAnswerState": true, "choices": ["$y_p(x) = ax+b$", "$y_p(x) = ax^2 + bx + c$&nbsp;", "$y_p(x) = ae^{\\simplify{({r1}x)}} + be^{\\simplify{({r2}x)}}$", "$y_p(x) = e^{ax}$"], "matrix": ["1", "0.5", 0, 0], "distractors": ["", "The $x^2$-term is unnecessary and will just make your life harder. But it's not technically wrong.", "This would make the LHS=0.", ""]}, {"type": "gapfill", "useCustomName": true, "customName": "Particular solution", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Solution of initial value problem", "rawLabel": "", "otherPart": 5, "variableReplacements": [], "availabilityCondition": "answered and credit=1", "penalty": "", "penaltyAmount": 0, "lockAfterLeaving": true}, {"label": "Try again!", "rawLabel": "Try again!", "otherPart": 4, "variableReplacements": [], "availabilityCondition": "", "penalty": "Retried particular solution", "penaltyAmount": "0.5", "lockAfterLeaving": false}, {"label": "Help, please!", "rawLabel": "Help, please!", "otherPart": 7, "variableReplacements": [], "availabilityCondition": "", "penalty": "", "penaltyAmount": 0, "lockAfterLeaving": false}, {"label": "Well, never mind. Let's press on!", "rawLabel": "Well, never mind. Let's press on!", "otherPart": 5, "variableReplacements": [], "availabilityCondition": "answered and credit<1", "penalty": "", "penaltyAmount": 0, "lockAfterLeaving": true}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Particular solution", "prompt": "<p>So, the particular solution is $y_p(x) = ax + b$.</p>\n<p>We compute the values of $a$ and $b$:</p>\n<p>$a = $&nbsp;[[0]]&nbsp; and</p>\n<p>$b = $&nbsp;[[1]]</p>\n<p><em>(If you get this wrong, you may retry this part and lose 0.5 marks. Thereafter you may retry this part as often as you like without losing further marks.)</em></p>", "gaps": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": "0.5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minValue": "a", "maxValue": "a", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": "0.5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minValue": "b", "maxValue": "b", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "Solution of initial value problem", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Try again!", "rawLabel": "Try again!", "otherPart": 5, "variableReplacements": [], "availabilityCondition": "answered and credit<1", "penalty": "Retried initial value problem", "penaltyAmount": "1", "lockAfterLeaving": false}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Solution to initial value problem", "prompt": "<p>So, we have determined that the general solution to the differential equation is</p>\n<p>\\[y(x) = C_1e^{\\simplify{({r1}x)}} +&nbsp;C_2e^{\\simplify{({r2}x)}} + \\simplify{({a}x+{b})}.\\]</p>\n<p>It remains to solve the initial value problem. Do so, and enter the resulting values of $C_1$ and $C_2$:</p>\n<p>$C_1 = $&nbsp;[[0]]</p>\n<p>$C_2 = $&nbsp;[[1]]</p>\n<p><em>(If you get this wrong, you may retry this part and lose 0.5 marks. Thereafter you may retry this part as often as you like without losing further marks.)</em></p>", "gaps": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": "0.5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "minValue": "c1", "maxValue": "c1", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": "0.5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Solution to initial value problem", "minValue": "c2", "maxValue": "c2", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "sortAnswers": false}, {"type": "information", "useCustomName": true, "customName": "Help with complementary function", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": true, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>We must first find the&nbsp;<em>complementary function</em>, that is the solution to the homogeneous version of this differential equation:</p>\n<p>\\[\\simplify{(y'' + {-r1-r2}y' + {r1*r2}y)} = 0.\\]</p>\n<p>We do this by setting up the&nbsp;<strong>characteristic equation</strong> of this differential equation (it comes from plugging $y(x)=e^{rx}$ into the differential equation and cancelling $e^{rx}$-terms):</p>\n<p>\\[\\simplify{(r^2 + {-r1-r2}r + {r1*r2})} = 0.\\]</p>\n<p>If this equation has two distinct solutions, $r=r_1$ and $r=r_2$, then the complementary function is</p>\n<p>\\[y_c(x) = C_1e^{r_1x} + C_2e^{r_2x}.\\]</p>\n<p>If the characteristic equation has a repeated real root $r$, then the complementary function is</p>\n<p>\\[y_c(x) = (C_1x+C_2)e^{rx}.\\]</p>\n<p>The third possibility is that the characteristic equation might have a conjugate pair of complex roots $r = \\alpha \\pm \\beta i$. In this case, the complementary function is&nbsp;</p>\n<p>\\[y_c(x) = e^{\\alpha x}\\big(C_1\\cos(\\beta x) + C_2\\sin(\\beta x)\\big).\\]</p>\n<p></p>"}, {"type": "information", "useCustomName": true, "customName": "Help with particular solution", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": true, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>We must now find a&nbsp;<em>particular solution</em>, i.e. a function that is a solution to our non-homogeneous differential equation.&nbsp;</p>\n<p>Think of it as guessing a function $y_p(x)$ such that, when plugged into $\\simplify{(y'' +{-r1-r2}y'+{r1*r2}y)}$ gives $\\simplify{({c}x+{d})}$.</p>\n<p>Our best bet is to try something of the same form as $\\simplify{({c}x+{d})}$, but with variable coefficients.</p>\n<p><strong>Warning:</strong> This won't work if any term of our $y_p(x)$ happens to make the LHS = 0, i.e. if any term of $y_p(x)$ has the same form as a term of the complementary function. If this happens, we need to multiply our $y_p(x)$ by $x$ and try again. For this reason, we should first find the complementary function before finding a particular solution.</p>"}], "partsMode": "explore", "maxMarks": "5", "objectives": [{"name": "Complementary function", "limit": "2", "mode": "sum"}, {"name": "Particular solution", "limit": "2", "mode": "sum"}, {"name": "Solution to initial value problem", "limit": "1", "mode": "sum"}, {"name": "Off to a good start (bonus point)", "limit": "1", "mode": "sum"}], "penalties": [{"name": "Retried complementary function", "limit": "0.5", "mode": "sum"}, {"name": "Retried particular solution", "limit": "0.5", "mode": "sum"}, {"name": "Retried initial value problem", "limit": "0.5", "mode": "sum"}], "objectiveVisibility": "when-active", "penaltyVisibility": "when-active", "contributors": [{"name": "Florian Breuer", "profile_url": "https://editor.newnumbas.net/accounts/profile/28/"}]}]}], "contributors": [{"name": "Florian Breuer", "profile_url": "https://editor.newnumbas.net/accounts/profile/28/"}]}